<img align="center" src="figures/course.png" width="800">

#                                    16720 (B) Neural Networks for Recognition - Assignment 3

     Instructor: Kris Kitani                       TAs: Qichen(Lead), Paritosh, Rawal, Yan, Zen, Wen-Hsuan

## Q2 Implement a Fully Connected Network (65 points + 10 Extra Credit)

**Please include all the write up answers below to theory.ipynb**

```python
# Do Not Modify
# Do Not Import ANY other packages
import numpy as np

# use for a "no activation" layer
def linear(x):
    return x

def linear_deriv(post_act):
    return np.ones_like(post_act)

def tanh(x):
    return np.tanh(x)

def tanh_deriv(post_act):
    return 1-post_act**2

def relu(x):
    return np.maximum(x, 0)

def relu_deriv(x):
    return (x > 0).astype(np.float)
```

### Q2.1 Network Initialization

#### Q2.1.1 (3 points WriteUp)
Why is it not a good idea to initialize a network with all zeros? If you imagine that every layer has weights and biases, what can a zero-initialized network output after training?

<font color="red">**Please include the write up answer to theory.ipynb**</font>

#### Q2.1.2 (3 points Autograder)
Implement `initialize_weights` below to initialize neural network weights with [Xavier initialization](http://proceedings.mlr.press/v9/glorot10a/glorot10a.pdf), where $Var[w] = \frac{2}{n_{in}+ n_{out}} $ where $n$ is the dimensionality of the vectors and you use a **uniform distribution** to sample random numbers (see eq 16 in [Xavier initialization](http://proceedings.mlr.press/v9/glorot10a/glorot10a.pdf)).

<!-- grade_id: cell-04dbce39e616d009, points: 0, grade: false, solution: true -->
```python
def initialize_weights(in_size: int, out_size: int, params: dict, name: str='' ):
    '''
    Initialize the weights W and b for a linear layer Y = XW + b
    
    [input]
    * in_size -- the feature dimension of the input
    * out_size -- the feature dimension of the output
    * params -- a dictionary containing parameters
    * name -- name of the layer
    
    HINTS:
    (1) b should be a 1D array, not a 2D array with a singleton dimension
    '''
    W, b = None, None

    # YOUR CODE HERE

    max_norm_w = np.sqrt(6)/np.sqrt(in_size+out_size)
    W = np.random.uniform(-max_norm_w, max_norm_w, (in_size, out_size))
    b = np.zeros(out_size)

    # raise NotImplementedError()

    params['W' + name] = W
    params['b' + name] = b
```

<!-- grade_id: q2_1_2, points: 3, grade: true, solution: false -->
```python
params = {}
initialize_weights(2,25,params,'layer1')
initialize_weights(25,4,params,'output')
assert(params['Wlayer1'].shape == (2,25))
assert(params['blayer1'].shape == (25,))

```

#### Q2.1.3 (2 points WriteUp)
Why do we initialize with random numbers? Why do we scale the initialization depending on layer size (see near Fig 6 in [Xavier initialization](http://proceedings.mlr.press/v9/glorot10a/glorot10a.pdf))?

<font color="red">**Please include the write up answer to theory.ipynb**</font>

### Q2.2 Forward Propagation

The appendix (in `theory.jpynb`) has the math for forward propagation, we will implement it here.

#### Q2.2.1 (12 points Autograder)
Implement `sigmoid`, along with `forward` propagation for a single layer with an activation function, namely
$y = \sigma(X W + b)$, returning the output and intermediate results for an $N \times D$ dimension input $X$, with examples along the rows, data dimensions along the columns.

<!-- grade_id: cell-9d629c29b34231c5, points: 0, grade: false, solution: true -->
```python
def sigmoid(X: np.ndarray):
    '''
    A sigmoid activation function
    
    [input]
    * X -- input data [N x D]
    
    [output]
    * res -- output after the sigmoid function
    '''
    res = None

    # YOUR CODE HERE

    res = 1 / (1+np.exp(-X.astype(np.float128)))

    # raise NotImplementedError()

    return res

```

<!-- grade_id: q2_2_1_sigmoid, points: 3, grade: true, solution: false -->
```python
test = sigmoid(np.array([-100,100]))
assert test.min() < 1e-3
assert test.max() > 1 - 1e-3

```

<!-- grade_id: cell-8a7b5e6897a9d712, points: 0, grade: false, solution: true -->
```python
def forward(X: np.ndarray, params: dict, name: str='',
            activation: callable=sigmoid):
    """
    Do a forward pass

    [input]
    * X -- input data [N x D]
    * params -- a dictionary containing parameters
    * name -- name of the layer
    * activation -- the activation function (default is sigmoid)
    
    [output]
    * post_act -- output after a linear layer and activation
    """
    pre_act, post_act = None, None
    # get the layer parameters
    W = params['W' + name]
    b = params['b' + name]

    # YOUR CODE HERE

    pre_act = np.dot(X, W) + b
    post_act = activation(pre_act)

    # raise NotImplementedError()

    # store the pre-activation and post-activation values
    # these will be important in backprop
    params['cache_' + name] = (X, pre_act, post_act)

    return post_act
```

<!-- grade_id: q2_2_1_forward, points: 9, grade: true, solution: false -->
```python
params = {'Wlayer1': np.random.rand(10, 25), 'blayer1': np.random.rand(25,)}
X = np.random.rand(3, 10)
y = forward(X, params, 'layer1')
assert 'cache_layer1' in params

```

#### Q2.2.2 (5 points Autograder)
Implement the `softmax` function. Be sure the use the numerical stability trick you derived in Q1.1 softmax.

<!-- grade_id: cell-e41cfc344ff8bdef, points: 0, grade: false, solution: true -->
```python
def softmax(X: np.ndarray):
    """
    A softmax function.
    
    [input]
    * X -- input data [N x D]
    
    [output]
    * res -- values after softmax
    """
    res = None

    # YOUR CODE HERE

    x_exp = np.exp(X-np.max(X,axis=1).reshape(-1, 1))
    exp_sum = np.sum(x_exp, axis=1).reshape(-1, 1)

    res = x_exp/exp_sum 

    #raise NotImplementedError()

    return res

```

<!-- grade_id: q2_2_2, points: 5, grade: true, solution: false -->
```python

```

#### Q2.2.3 (5 points Autograder)
Implement `compute_loss_and_acc` to compute the accuracy of a set of labels, along with the scalar loss across the data.  The loss function generally used for classification is the cross-entropy loss.

$$L_{\textbf{f}}(\textbf{D}) = - \sum_{(\textbf{x}, \textbf{y})\in \textbf{D}}\textbf{y}\cdot\log(\textbf{f}(\textbf{x}))$$
Here $\textbf{D}$ is the full training dataset of data samples $\textbf{x}$ ($N\times 1$ vectors, N = dimensionality of data) and labels $\textbf{y}$ ($C\times 1$ one-hot vectors, C = number of classes).

<!-- grade_id: cell-e8d3b0354b03c98c, points: 0, grade: false, solution: true -->
```python
def compute_loss_and_acc(y: np.ndarray, probs: np.ndarray):
    """
    Compute total loss and accuracy
    
    [input]
    * y -- one hot labels [N x C]
    * probs -- class probabities [N x C]
    
    [output]
    * loss -- cross-entropy loss
    * acc -- accuracy
    """
    loss, acc = None, None

    # YOUR CODE HERE

    N = y.shape[0]
    loss = -np.sum(y*np.log(probs))
    acc = np.sum(np.argmax(probs, axis=1) == np.argmax(y, axis=1))/N

    # raise NotImplementedError()

    return loss, acc
```

<!-- grade_id: q2_2_3, points: 5, grade: true, solution: false -->
```python

```

### Q2.3 Backwards Propagation

#### Q2.3.1 (10 points Autograder)
Compute back-propagation for a single layer, given the original weights, the appropriate intermediate results, and given gradient with respect to the loss. You should return the gradient with respect to $X$ so you can feed it into the next layer. As a size check, your gradients should be the same dimensions as the original objects.

<!-- grade_id: cell-05a2aee350a000f3, points: 0, grade: false, solution: true -->
```python
def sigmoid_deriv(post_act: np.ndarray):
    """
    Derivative of sigmoid.
    
    we give this to you because you proved it
    it's a function of post_act
    """
    res = post_act*(1.0-post_act)
    return res


def backwards(delta: np.ndarray, params: dict, name: str='',
              activation_deriv: callable=sigmoid_deriv):
    """
    Do a backwards pass

    [input]
    * delta -- errors to backprop
    * params -- a dictionary containing parameters
    * name -- name of the layer
    * activation_deriv -- the derivative of the activation_func
    
    [output]
    * grad_X -- gradient w.r.t X
    """
    grad_X, grad_W, grad_b = None, None, None
    # everything you may need for this layer
    W = params['W' + name]
    b = params['b' + name]
    X, pre_act, post_act = params['cache_' + name]

    # do the derivative through activation first
    # then compute the derivative W,b, and X
    # YOUR CODE HERE

    deriv = activation_deriv(post_act) * delta
    grad_W = np.dot(np.transpose(X), deriv)
    grad_b = np.sum(deriv, axis=0)
    grad_X = np.dot(deriv, np.transpose(W))

    # raise NotImplementedError()

    # store the gradients
    params['grad_W' + name] = grad_W
    params['grad_b' + name] = grad_b
    return grad_X
```

<!-- grade_id: q2_3_1, points: 10, grade: true, solution: false -->
```python
# we use random values to test your implementation 
# independent of previous questions
n, c1, c2 = 5, 40, 20 
delta = np.random.rand(n, c2)
name = 'layer1'
params = {
    'W'+name: np.random.rand(c1, c2),
    'b'+name: np.random.rand(c2),
    'cache_'+name: (np.random.rand(n, c1), 
                     np.random.rand(n, c2), 
                     np.random.rand(n, c2))
}

grad = backwards(delta, params, name, linear_deriv)

assert 'grad_W' + name in params
assert 'grad_b' + name in params

assert params['grad_W'+name].shape == params['W'+name].shape
assert params['grad_b'+name].shape == params['b'+name].shape

```

### Q2.4 Convolutional Layer [Extra Credit]
**Note: We would recommend finishing other questions before attempting questions under this section (q2.4.1 and q2.4.2)**

For now we have worked with linear layer in fully-connected networks. In practice, convolutional layers are commonly used to extract image feature. You will implement the forward and backawad propagation for convolutional layer in this subsection. 

#### Q2.4.1 [Extra Credit](5 points Autograder)
Similar to Q2.2.1, implement `conv_forward` for a single convolutional layer with zero paddings.

<!-- grade_id: cell-5e3595068856df10, points: 0, grade: false, solution: true -->
```python
'''

def conv_forward(X: np.ndarray, params: dict, name: str='',
            stride: int=1, pad: int=0):
    """
    Do a forward pass for a convolutional layer

    [input]
    * X -- input data [N x C x H x W]
    * params -- a dictionary containing parameters
    * name -- name of the layer
    * stride, pad -- convolution parameters
    
    [output]
    * res -- output after a convolutional layer
    """
    res = None
    # get the layer parameters
    w = params['W' + name] # Conv Filter weights [F x C x HH x WW]
    b = params['b' + name] # Biases [F]
    
    # YOUR CODE HERE

    

    # raise NotImplementedError()

    # store the input and convolution parameters
    # these will be important in backprop
    params['cache_' + name] = (X, stride, pad)

    return res

'''
```

Output:
```
'\n\ndef conv_forward(X: np.ndarray, params: dict, name: str=\'\',\n            stride: int=1, pad: int=0):\n    """\n    Do a forward pass for a convolutional layer\n\n    [input]\n    * X -- input data [N x C x H x W]\n    * params -- a dictionary containing parameters\n    * name -- name of the layer\n    * stride, pad -- convolution parameters\n    \n    [output]\n    * res -- output after a convolutional layer\n    """\n    res = None\n    # get the layer parameters\n    w = params[\'W\' + name] # Conv Filter weights [F x C x HH x WW]\n    b = params[\'b\' + name] # Biases [F]\n    \n    # YOUR CODE HERE\n\n    \n\n    # raise NotImplementedError()\n\n    # store the input and convolution parameters\n    # these will be important in backprop\n    params[\'cache_\' + name] = (X, stride, pad)\n\n    return res\n\n'
```

<!-- grade_id: q2_4_1, points: 5, grade: true, solution: false -->
```python
'''

x_shape = np.array((2, 3, 4, 4))
w_shape = np.array((3, 3, 4, 4))
x = np.linspace(-0.1, 0.5, num=np.prod(x_shape), dtype=np.float64).reshape(*x_shape)
w = np.linspace(-0.2, 0.3, num=np.prod(w_shape), dtype=np.float64).reshape(*w_shape)
b = np.linspace(-0.1, 0.2, num=3, dtype=np.float64)

params = {'WConv_layer1': w, 'bConv_layer1': b}
y = conv_forward(np.array(x), params, 'Conv_layer1', stride=2, pad=1)
assert 'cache_Conv_layer1' in params


y_ref = np.array([[[[-0.08759809, -0.10987781],
                              [-0.18387192, -0.2109216 ]],
                             [[ 0.21027089,  0.21661097],
                              [ 0.22847626,  0.23004637]],
                             [[ 0.50813986,  0.54309974],
                              [ 0.64082444,  0.67101435]]],
                            [[[-0.98053589, -1.03143541],
                              [-1.19128892, -1.24695841]],
                             [[ 0.69108355,  0.66880383],
                              [ 0.59480972,  0.56776003]],
                             [[ 2.36270298,  2.36904306],
                              [ 2.38090835,  2.38247847]]]], 
            )
assert y.shape == y_ref.shape

max_diff = np.max(np.abs((y_ref - y)))
base = (np.abs(y_ref) + np.abs(y)).clip(np.finfo(float).eps).max()
print(max_diff/base) # the difference should be less than 1e-8

'''
```

Output:
```
"\n\nx_shape = np.array((2, 3, 4, 4))\nw_shape = np.array((3, 3, 4, 4))\nx = np.linspace(-0.1, 0.5, num=np.prod(x_shape), dtype=np.float64).reshape(*x_shape)\nw = np.linspace(-0.2, 0.3, num=np.prod(w_shape), dtype=np.float64).reshape(*w_shape)\nb = np.linspace(-0.1, 0.2, num=3, dtype=np.float64)\n\nparams = {'WConv_layer1': w, 'bConv_layer1': b}\ny = conv_forward(np.array(x), params, 'Conv_layer1', stride=2, pad=1)\nassert 'cache_Conv_layer1' in params\n\n\ny_ref = np.array([[[[-0.08759809, -0.10987781],\n                              [-0.18387192, -0.2109216 ]],\n                             [[ 0.21027089,  0.21661097],\n                              [ 0.22847626,  0.23004637]],\n                             [[ 0.50813986,  0.54309974],\n                              [ 0.64082444,  0.67101435]]],\n                            [[[-0.98053589, -1.03143541],\n                              [-1.19128892, -1.24695841]],\n                             [[ 0.69108355,  0.66880383],\n            
```

#### Q2.4.2 [Extra Credit](5 points Autograder)
Implement `conv_backword` for a single convolutional layer with zero paddings.
Compute back-propagation for a single convolutional layer, given the original weights, the cached input, and given gradient with respect to the loss. Similar to Q2.3.1, you should return the gradient with respect to $X$ so you can feed it into the next layer. As a size check, your gradients should be the same dimensions as the original objects.

<!-- grade_id: cell-cc53a61e4f74a964, points: 0, grade: false, solution: true -->
```python
'''

def conv_backward(delta: np.ndarray, params: dict, name: str=''):
    """
    Do a backwards pass for a convolutional layer

    [input]
    * delta -- errors to backprop
    * params -- a dictionary containing parameters
    * name -- name of the layer
    
    [output]
    * grad_X -- gradient w.r.t X
    """
    grad_X, grad_W, grad_b = None, None, None
    # everything you may need for this layer
    W = params['W' + name]
    b = params['b' + name]
    X, stride, pad = params['cache_' + name]

    # compute the derivative W,b, and X
    # YOUR CODE HERE
    raise NotImplementedError()

    # store the gradients
    params['grad_W' + name] = grad_W
    params['grad_b' + name] = grad_b
    return grad_X

'''
```

Output:
```
'\n\ndef conv_backward(delta: np.ndarray, params: dict, name: str=\'\'):\n    """\n    Do a backwards pass for a convolutional layer\n\n    [input]\n    * delta -- errors to backprop\n    * params -- a dictionary containing parameters\n    * name -- name of the layer\n    \n    [output]\n    * grad_X -- gradient w.r.t X\n    """\n    grad_X, grad_W, grad_b = None, None, None\n    # everything you may need for this layer\n    W = params[\'W\' + name]\n    b = params[\'b\' + name]\n    X, stride, pad = params[\'cache_\' + name]\n\n    # compute the derivative W,b, and X\n    # YOUR CODE HERE\n    raise NotImplementedError()\n\n    # store the gradients\n    params[\'grad_W\' + name] = grad_W\n    params[\'grad_b\' + name] = grad_b\n    return grad_X\n\n'
```

<!-- grade_id: q2_4_2, points: 5, grade: true, solution: false -->
```python
'''

x = np.random.rand(5, 4, 16, 16)
w = np.random.rand(8, 4, 7, 7)
b = np.random.rand(8,)
dout = np.random.rand(5, 8, 16, 16)

params = {'WConv_layer1': w, 'bConv_layer1': b}
y = conv_forward(x, params, 'Conv_layer1', stride=2, pad=3)
dx = conv_backward(dout, params, 'Conv_layer1')

assert x.shape == dx.shape
assert params['grad_WConv_layer1'].shape == params['WConv_layer1'].shape
assert params['grad_bConv_layer1'].shape == params['bConv_layer1'].shape

'''
```

Output:
```
"\n\nx = np.random.rand(5, 4, 16, 16)\nw = np.random.rand(8, 4, 7, 7)\nb = np.random.rand(8,)\ndout = np.random.rand(5, 8, 16, 16)\n\nparams = {'WConv_layer1': w, 'bConv_layer1': b}\ny = conv_forward(x, params, 'Conv_layer1', stride=2, pad=3)\ndx = conv_backward(dout, params, 'Conv_layer1')\n\nassert x.shape == dx.shape\nassert params['grad_WConv_layer1'].shape == params['WConv_layer1'].shape\nassert params['grad_bConv_layer1'].shape == params['bConv_layer1'].shape\n\n"
```

### Q2.5 Training Loop
You will tend to see gradient descent in three forms: "normal", "stochastic" and "batch". "Normal" gradient descent aggregates the updates for the entire dataset before changing the weights. Stochastic gradient descent applies updates after every single data example. Batch gradient descent is a compromise, where random subsets of the full dataset are evaluated before applying the gradient update. 

#### Q2.5.1 (10 points Autograder)
Write a training loop that generates random batches, iterates over them for many iterations, does forward and backward propagation, and applies a gradient update step. Specifically, implement `get_random_batches` and `train` functions below.

<!-- grade_id: cell-1b13b59b8a93b801, points: 0, grade: false, solution: true -->
```python
def get_random_batches(x: np.ndarray, y: np.ndarray, batch_size: int) -> list:
    """
    Split x and y into random batches
    
    [input]
    * x -- training samples
    * y -- training lables
    * batch_size -- batch size
    
    [output]
    * batches -- a list of [(batch1_x,batch1_y)...]
    """
    # 
    # return 
    batches = []
    # YOUR CODE HERE

    num = x.shape[0]
    idx = np.arange(0, num)
    np.random.shuffle(idx)
    shuffled_x = x[idx]
    shuffled_y = y[idx]
    for i in range(0, num, batch_size):
        last_idx = min(num, i+batch_size)
        batches.append((shuffled_x[i: last_idx], shuffled_y[i: last_idx]))

    # raise NotImplementedError()
    return batches

```

<!-- grade_id: q2_5_1_random_batches, points: 2, grade: true, solution: false -->
```python
n, c1, c2 = 20, 100, 5
batch_size = 3
x = np.random.rand(n, c1)
y = np.random.rand(n, c2)
batches = get_random_batches(x, y, batch_size)
assert type(batches) == list
assert len(batches) >= 6

```

<!-- grade_id: cell-96bdeec247659603, points: 0, grade: false, solution: true -->
```python
def train(x: np.ndarray, y: np.ndarray, params: dict, batch_size: int = 5,
          max_iters: int = 500, learning_rate: float=1e-3):
    
    """
    Train the network with two sequential layers: 
    (1) one layer named "layer1" with sigmoid activation
    (2) one layer named "output" with softmax activation

    [input]
    * x -- training samples
    * y -- training lables
    * params -- a dictionary containing initial parameters
    * batch_size -- batch size
    * max_iters -- total number of iterations
    * learning_rate -- learning rate
    
    [output]
    * total_loss, avg_acc -- loss and accuracy for the last iteration
    """

    batches = get_random_batches(x, y, batch_size)
    # print(len(batches))

    for itr in range(max_iters):
        total_loss = 0
        avg_acc = 0
        for xb, yb in batches:

            # forward
            # YOUR CODE HERE

            h1 = forward(xb, params, 'layer1')
            probs = forward(h1, params, 'output', softmax)

            # raise NotImplementedError()
            
            # loss
            # be sure to add loss and accuracy to epoch totals
            # YOUR CODE HERE

            loss, acc = compute_loss_and_acc(yb, probs)
            total_loss += loss
            avg_acc += acc*xb.shape[0]/x.shape[0]

            # raise NotImplementedError()
            
            # backward
            # YOUR CODE HERE

            delta1 = probs
            probs_len = probs.shape[0]
            delta1[np.arange(probs_len), np.argmax(yb, axis=1)] -= 1
            delta2 = backwards(delta1, params, 'output', linear_deriv)
            backwards(delta2, params, 'layer1', sigmoid_deriv)

            # raise NotImplementedError()

            # apply gradient
            # YOUR CODE HERE

            params['Wlayer1'] -= learning_rate*params['grad_Wlayer1']
            params['blayer1'] -= learning_rate*params['grad_blayer1']
            params['Woutput'] -= learning_rate*params['grad_Woutput']
            params['boutput'] -= learning_rate*params['grad_boutput']

            # raise NotImplementedError()
            
        if itr % 100 == 0:
            print("itr: {:02d} \t loss: {:.2f} \t acc : {:.2f}".format(
                itr, total_loss, avg_acc))
    return total_loss, avg_acc

```

<!-- grade_id: q2_5_1_train, points: 8, grade: true, solution: false -->
```python
# Successulf implementation of dependent functions are required to get full score for the `train` function

# create inputs
g0 = np.random.multivariate_normal([3.6,40],[[0.05,0],[0,10]],10)
g1 = np.random.multivariate_normal([3.9,10],[[0.01,0],[0,5]],10)
g2 = np.random.multivariate_normal([3.4,30],[[0.25,0],[0,5]],10)
g3 = np.random.multivariate_normal([2.0,10],[[0.5,0],[0,10]],10)
x = np.vstack([g0,g1,g2,g3])

# create labels
y_idx = np.array([0 for _ in range(10)] + [1 for _ in range(10)] + [2 for _ in range(10)] + [3 for _ in range(10)])

# turn labels to one_hot
y = np.zeros((y_idx.shape[0],y_idx.max()+1))
y[np.arange(y_idx.shape[0]),y_idx] = 1

# parameters in a dictionary
params = {}
# initialize a layer
initialize_weights(2,25,params,'layer1')
initialize_weights(25,4,params,'output')

# print(x.shape)
# print(y.shape)

# train the two-layer neural network
total_loss, avg_acc = train(x, y, params, batch_size=5, max_iters=500, learning_rate=1e-3)
print("itr: {:02d} \t loss: {:.2f} \t acc : {:.2f}".format(500, total_loss, avg_acc))

# with default settings, you should get loss < 35 and accuracy > 75%
assert total_loss < 35 and avg_acc > 0.75

```

Output:
```
itr: 00 	 loss: 66.44 	 acc : 0.25
itr: 100 	 loss: 42.29 	 acc : 0.72
itr: 200 	 loss: 35.68 	 acc : 0.80
itr: 300 	 loss: 32.02 	 acc : 0.80
itr: 400 	 loss: 29.51 	 acc : 0.85
itr: 500 	 loss: 27.55 	 acc : 0.85
```

### Q2.6 Numerical Gradient Checker

#### Q2.6.1 (15 points Autograder)
Implement the `centeral_differences_gradient` function. Instead of using the analytical gradients computed from the chain rule, add $\epsilon$ offset to each element in the weights, and compute the numerical gradient of the loss with central differences. Central differences is just $\frac{f(x+\epsilon) - f(x-\epsilon)}{2 \epsilon}$. Remember, this needs to be done for each scalar dimension in all of your weights independently. 

<!-- grade_id: cell-2e7c04b3366231a7, points: 0, grade: false, solution: true -->
```python
def centeral_differences_gradient(params: dict, eps = 1e-6):
    """
    Compute the estimated gradients using central difference
    
    Hint:
    please feel free to reuse the functions above
    """
    for k, v in params.items():
        if '_' in k:
            continue
        # we have a real parameter!
        # for each value inside the parameter
        #   subtract epsilon
        #   run the forward function by 
        #   get the loss
        #   add epsilon
        #   run the forward function
        #   get the loss
        #   compute derivative with central diffs
    
        # YOUR CODE HERE

        if (k == 'cache_layer1'):
            params = copy.deepcopy(params_orig)
            X, pre_act, post_act = v

            h, w = params["Wlayer1"].shape
            for i in range(h):
                for j in range(w):
                    params["Wlayer1"][i][j] = params_orig["Wlayer1"][i][j]+eps
                    h1 = forward(X, params, 'layer1')
                    probs = forward(h1, params, 'output', softmax)
                    loss1 = -np.sum(y*np.log(probs))

                    params["Wlayer1"][i][j] = params_orig["Wlayer1"][i][j]-eps
                    h1 = forward(X, params, 'layer1')
                    probs = forward(h1, params, 'output', softmax)
                    loss2 = -np.sum(y*np.log(probs))

                    params["Wlayer1"][i][j] = params_orig["Wlayer1"][i][j]
                    params["grad_Wlayer1"][i][j] = (loss1-loss2)/(2*eps)

            h = params["blayer1"].shape[0]
            for i in range(h):
                params["blayer1"][i] = params_orig["blayer1"][i][j]+eps
                h1 = forward(X, params, 'layer1')
                probs = forward(h1, params, 'output', softmax)
                loss1 = -np.sum(y*np.log(probs))

                params["blayer1"][i] = params_orig["blayer1"][i][j]-eps
                h1 = forward(X, params, 'layer1')
                probs = forward(h1, params, 'output', softmax)
                loss2 = -np.sum(y*np.log(probs))

                params["blayer1"][i] = params_orig["blayer1"][i]
                params["grad_blayer1"][i] = (loss1-loss2)/(2*eps)
                
        if (k == 'cache_ouput'):
            X, pre_act, post_act = v

            h, w = params["Woutput"].shape
            for i in range(h):
                for j in range(w):
                    params["Woutput"][i][j] = params_orig["Woutput"][i][j]+eps
                    probs = forward(X, params, 'output', softmax)
                    loss1 = -np.sum(y*np.log(probs))

                    params["Woutput"][i][j] = params_orig["Woutput"][i][j]-eps
                    probs = forward(X, params, 'output', softmax)
                    loss2 = -np.sum(y*np.log(probs))

                    params["Woutput"][i][j] = params_orig["Woutput"][i][j]
                    params["grad_Woutput"][i][j] = (loss1-loss2)/(2*eps)

            h = params["boutput"].shape[0]
            for i in range(h):
                params["boutput"][i] = params_orig["boutput"][i][j]+eps
                probs = forward(X, params, 'output', softmax)
                loss1 = -np.sum(y*np.log(probs))

                params["boutput"][i] = params_orig["boutput"][i][j]-eps
                probs = forward(h1, params, 'output', softmax)
                loss2 = -np.sum(y*np.log(probs))

                params["boutput"][i] = params_orig["boutput"][i]
                params["grad_boutput"][i] = (loss1-loss2)/(2*eps)
        # raise NotImplementedError()

```

<!-- grade_id: q2_6_1, points: 15, grade: true, solution: false -->
```python
# Compute the analytical gradients
h1 = forward(x,params,'layer1')
probs = forward(h1,params,'output',softmax)
delta1 = probs
delta1[np.arange(probs.shape[0]),y_idx] -= 1

delta2 = backwards(delta1,params,'output',linear_deriv)
backwards(delta2,params,'layer1',sigmoid_deriv)

import copy
params_orig = copy.deepcopy(params)

# Compute the estimated gradient using central difference
centeral_differences_gradient(params)

total_error = 0
for k in params.keys():
    if 'grad_' in k:
        # relative error
        err = np.abs(params[k] - params_orig[k])/np.maximum(np.abs(params[k]),np.abs(params_orig[k]))
        err = err.sum()
        total_error += err
# should be less than 1e-4

print(total_error)
assert 0. <= total_error < 1e-4
```

Output:
```
0.0
```